# Homework 2

Submit your assignment as a Jupyter notebook file called `yourmacid_hw2.ipynb`, where `yourmacid` is replaced with your MacID from your McMaster email address. (For example: since my McMaster email address is bolker@mcmaster.ca, I would submit a file called `bolker_hw2.ipynb`.) Submit this file to the Dropbox on the MATH 1MP3 Avenue to Learn site by **Weds 2 October at 11:59 PM**.

To complete your assignment, edit the function definitions in this template (then rename the file appropriately, as above). (Don't forget to use appropriate `return()` statements in your functions!) **Do not alter any other part of this template.**

While developing and testing your code, it might be helpful to use
`print` statements, but please remove them before submitting your solution. The template contains several `assert` statements that when executed will test your code.

Any file submitted for grading that does not conform to the specifications above will lead to a grade of zero on the assignment.

Before submitting your solution, you should make sure that it runs properly. To do this, you can open it in the Jupyter notebook server that comes installed with Anaconda, or you can use the McMaster Jupyter server, by going to the website https://mcmaster.syzygy.ca/. (Try running `Kernel > Restart & Aun all` from the menu to run all of your code from scratch and make sure nothing is left out.) Feel free to use the Spyder IDE, or some other IDE, to develop your code before entering it into the Jupyter notebook that you submit.

Your grade for each question will depend on whether or not your code correctly handles not just the test cases provided in the template, but other test cases that will be run on your solutions.

Do not leave this assignment until the last minute; last-minute computer/internet/Avenue issues on your part are your problem, not ours ...

Late assignments will be penalized 25% per day.

All work submitted for grading must be your own. You may discuss homework problems and related material with other students, but you must not submit work copied from others or from the internet.  If you use any internet resources for help, you **must** list them as comments (##) in your answer.

You should not `import` any modules for this homework, other than the `math` module (you may lose points if you do). 

```python
## run this code chunk before using the rest of the notebook ...
import math
```

## 1. Coin counting (5 points)

Given an amount of money and a list of denominations, write a function that returns a list of integers (`int`) that gives the (smallest) number of coins from a given combination of units; if the remaining number of cents (an integer) is not equal to zero, append it to the result. For example, Your function should return `[2,1,2,2,0,3]` for arguments `total = 5.73` and `denoms = [2,1,0.25,0.1,0.05]`. Do not assume that the list of denominations is correctly sorted; use `help("list.sort")` to find out about the sorting method for lists and its `reverse` argument (to sort the list into reverse order).

To get the remaining number of cents, use `round(x*100)` where `x` is the remaining amount of value (e.g. if `x` is 0.02, this will give a value of 2).

<!-- grade_id: cell-6c5310447fca2b1e, points: 0, grade: false, solution: true -->
```python
def coin_counting(total, denoms):
    """Compute the smallest number of coins for a 
    given combination of units, and the remaining number
    of cents (if any)"""
    # YOUR CODE HERE
    raise NotImplementedError()
```

Your function should return `[2,1,2,2,0,3]` for arguments `total = 5.73`, `denoms = [2,1,0.25,0.1,0.05]`, and `[2,1,0,0,0]` for arguments `total = 5`, `denoms = [2,1,0.25,0.1,0.05]`.
Check that it does:

```python
coin_counting(5.73,[2,1,0.25,0.1,0.05])
```

```python
coin_counting(5,[2,1,0.25,0.1,0.05])
coin_counting(3.61,[2,1,0.25,0.05])
```

The `assert` keyword is a way to check that your function acts the way it's supposed to: if the statement after `assert` is not `True`, Python will give an error.

<!-- grade_id: cell-b8da04c26c34814c, points: 2, grade: true, solution: false -->
```python
## Check that coin_counting returns the correct output for several inputs"
assert coin_counting(5.73,[2,1,0.25,0.1,0.05]) == [2,1,2,2,0,3]

```

<!-- grade_id: cell-26a347c2f6954569, points: 3, grade: true, solution: false -->
```python
assert coin_counting(5,[2,1,0.25,0.1,0.05]) == [2,1,0,0,0]
```

## Question 2. Matrix row/column sums (5 points)

Using nested `for` loops, write a function that returns a list of the sums by rows and columns for a given matrix. For example, your function should return `[[6,15,24],[12,15,18]]` for a matrix argument `[[1,2,3],[4,5,6],[7,8,9]]`.

<!-- grade_id: cell-dfc921922f1df1f1, points: 0, grade: false, solution: true -->
```python
def sum_row_column(mat):
    """Compute the sum of rows and columns for 
    a matrix specified as a list of lists of numbers"""
    # YOUR CODE HERE
    raise NotImplementedError()
```

Your function should return `[[6,15,24],[12,15,18]]` for a given matrix `[[1,2,3],[4,5,6],[7,8,9]]`. 
Check that it does:

```python
sum_row_column([[1,2,3],[4,5,6],[7,8,9]])
```

<!-- grade_id: cell-3b3386ed7511f5d6, points: 1, grade: true, solution: false -->
```python
## Check that sum_row_column returns the correct answer for various inputs."""
assert sum_row_column([[1,2,3],[4,5,6],[7,8,9]]) == [[6,15,24],[12,15,18]]
```

<!-- grade_id: cell-b0940909d03bcb43, points: 3, grade: true, solution: false -->
```python
assert sum_row_column([[1]]) == [[1],[1]]
```

<!-- grade_id: cell-063225cc742b89b0, points: 1, grade: true, solution: false -->
```python
## ignore this cell please
```

## 3. Standard deviation of a numeric sample (5 points)

The sample standard deviation of a set of numbers is defined as

$$
s = \sqrt{\frac{\sum_{i=1}^n (x_i-\bar x)^2}{n-1}}
$$
where $\bar x$ is the mean ($(\sum x)/n$) (the sum here starts from 1 because that's standard math notation, but you should use 0-indexing in your answer!)  Write a Python function `stddev(x)` that computes the standard deviation of a list `x` containing numeric elements.

The easiest way to do this is to write two loops, the first one of which calculates $\sum x$ (so that you can calculate $\bar x$ to use in the second loop).

* You'll need to use `math.sqrt()` to compute the square root.
* There is another way to do this computation with a single `for` loop (by computing $\sum x$ and $\sum x^2$ in the same loop, and using some algebra); this is faster but may be *slightly* less accurate.

<!-- grade_id: cell-3929728463c15cc7, points: 0, grade: false, solution: true -->
```python
def stddev(x):
    """compute the standard deviation of a list of numbers x"""
    # YOUR CODE HERE
    raise NotImplementedError()
```

A good first test of your function is that the standard deviation of a homogeneous list (i.e., where all elements are the same) should be 0.


<!-- grade_id: cell-56a5e67611dd7c3e, points: 1, grade: true, solution: false -->
```python
assert stddev([1,1,1,1,1]) == 0
```

`math.isclose()` is a utility function that tests to see whether two `float` values are *almost* equal (we will see in a few weeks why testing exact equality is usually a bad idea)

<!-- grade_id: cell-d8c814a1639b65b8, points: 1, grade: true, solution: false -->
```python
assert math.isclose(stddev([1,2,1,3,4]),1.3038404810405297)
```

Another good way to test code is to pick random numbers to test it with. We use `random.seed()` first to make sure that the values come out the same for everyone (we will cover this in more depth later ...)

<!-- grade_id: cell-052b1f1c224a2675, points: 1, grade: true, solution: false -->
```python
import random
random.seed(12345)  ## initialize random-number generator
x = []              ## initialize empty list
for i in range(20):
    x.append(random.randint(1,10))  ## pick a random int from 1 to 10
assert math.isclose(stddev(x), 2.704771612111494)
```

<!-- grade_id: cell-24e840fec7775e5b, points: 1, grade: true, solution: false -->
```python
## ignore this cell please
```

<!-- grade_id: cell-3cecc81e1b16ff34, points: 1, grade: true, solution: false -->
```python
## ignore this cell please
```


## 4. List/string manipulation (5 points)

Write a function that accepts a list of strings.  If the list is in the format `["givenname","familyname"]` (i.e. exactly two alphabetic strings), it should return a single string made up of two (first-letter) capitalized names, separated by a space; otherwise it should return `False`. For example, your function should return "Weijie Pang" for a given list `["weijie","Pang"]`; it should return `False` for arguments `["firstname"]` or `["Joe","2"]` or `["Joe","Q","Public"]`.

You should probably make use of the `.isalpha()` and `.capitalize()` string methods.


<!-- grade_id: cell-278142589173bda8, points: 0, grade: false, solution: true -->
```python
def cap_name(name):
    """Return a capitalized name for a valid name list, 
    or False otherwise"""
    # YOUR CODE HERE
    raise NotImplementedError()
```

Your function should return "Ben Bolker" for a given list `["ben","bolker"]`. It should return `False` for a given list `["weijie ", "pang"]` (because of the extra space in the first name). Check that it does:

```python
cap_name(["ben","bolker"])
```

```python
cap_name(["weijie ", "pang"])
```

<!-- grade_id: cell-a3c40f5510b60064, points: 2, grade: true, solution: false -->
```python
## Check that cap_name returns the correct output for several inputs
assert cap_name(["Ben","Bolker"]) == "Ben Bolker"

```

<!-- grade_id: cell-3d232f543f0a7f13, points: 3, grade: true, solution: false -->
```python
## Check that cap_name returns the correct output for several inputs
assert cap_name(["weijie ", "pang"]) == False
```

## 5. Taylor series (6 points)

The Taylor series of $\sin(x)$ around $x=0$ (also called the *Maclaurin series*) is
$$
x - x^3/3! + x^5/5! - x^7/7! + ...
$$
or
$$
\sum_{i=0}^\infty (-1)^i \frac{x^{2i+1}}{(2i+1)!}
$$
where $!$ stands for the factorial function, i.e. $j! = j \cdot (j-1) \cdot (j-2) \cdot \dots 1$
Write a function that computes this expression, adding up terms until the absolute value of the *next* term is less than a specified tolerance `tol` or until at most `nmax` terms have been added. 

Hints:

* the default tolerance value `1e-7` is expressed in ["e notation"](https://en.wikipedia.org/wiki/Scientific_notation#E-notation); in this notation `[x]e[y]` denotes $x \cdot 10^{y}$ (so `1e-7==0.0000001`)
* when translating math into computer code, remember that multiplication is usually denoted by *juxtaposition* in math notation (i.e. $ab$ is "$a$ times $b$") but that you need a multiplication operator (`*`) in Python (i.e. `a*b`)
* `abs()` is the absolute value function in Python
* it probably makes sense for you to initialize the value of the current term to `x` and the value of the factorial to 1 before starting your `while` loop; set `i=0`, then increment it by 1 (i.e. add 1 to it) at the beginning of your `while` loop code. Then you can keep track of the value of the factorial by multiplying `f` by `(2*i)*(2*i+1)` at each step.
* initialize `result` to 0; after your `while` loop finishes, check to see if `i` is still zero after the `while` loop finishes (i.e., the loop didn't run at all), return the value of `x` as the result.
* If you want to compute $-1^i$ in Python you need to use parentheses around the -1, `(-1)**i`. Or you could keep track of a sign variable (initialize it to 1 and multiply it by -1 at each step).
* use `print()` statements to debug your code.


<!-- grade_id: cell-48e495707cebea72, points: 0, grade: false, solution: true -->
```python
def sin_taylor(x, tol=1e-7, nmax=100):
    """compute the Taylor series of sin(x)
    around x=0 to a specifies tolerance/max number of terms"""
    # YOUR CODE HERE
    raise NotImplementedError()
```

Your function should return 0 for `x=0` (this input will cause an infinite loop if you're not testing for the maximum number of steps!) and a value close to 1 for `x=math.pi/2`

<!-- grade_id: cell-e8b84068612434b3, points: 2, grade: true, solution: false -->
```python
assert sin_taylor(0)==0
assert math.isclose(sin_taylor(math.pi/2),0.999999943741051)
```

You should also test that `sin_taylor(x,nmax=1)` is equal to `x`, and that `sin_taylor(x,tol=math.inf)` (`math.inf` is the Python expression for an infinite value) is also equal to `x` for any value of `x` ...

<!-- grade_id: cell-d085634bfcdf209e, points: 4, grade: true, solution: false -->
```python
assert sin_taylor(17,nmax=1)==17
### HIDDEN TESTS
assert sin_taylor(math.pi/2,tol=0.1)==0.9248322292886504
assert math.isclose(sin_taylor(2*math.pi),0,abs_tol=1e-6)
```